# Week 2 Programming Exercises

See: _McKinney 2.3_ and [Python Documentation](https://docs.python.org/3/tutorial/introduction.html) section 3.1.1 on using Python as a calculator.

**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 5.1 BMI Calculation

Create new names for the variables that you need in the code framework below and use the variables to compute the corresponding body mass index (BMI). You can use any meaningful variable names you like. When you compute the BMI, assign a variable to it, and then print the result to the notebook's output.

* Weght in kilograms is 90.7
* Height in meters in 1.82

Note that a simple BMI calculation is weight in kilograms divided by the square of height in meters.

$ bmi = \frac{weight\,\mathrm{kg}}{(height\,\mathrm{m})^2} $

For this problem, you'll be graded both on proper calculation as well as on programming style and appropriate variables names.

_Also note that BMI should not be considered a universal measure of health or fitness. A health BMI can vary greatly depending on many other fitness conditions and body types._

<!-- grade_id: bmi-answer, points: 0, grade: false, solution: true -->
```python
def bmi(weight_kg, height_m):

    # 1. Change the names of parameter1 and parameter2 to be something more meaningful
    # 2. Change the name of the return variable below

    ### BEGIN SOLUTION
    answer = -1
    ### END SOLUTION

    return answer
```

```python
def bmi(body_weight_in_kg, height_in_mts):
  # Body Mass Index(BMI) is calculated using the formula: weight kg/(height mts**2)
  bmi= body_weight_in_kg/(height_in_mts**2)
  # Return to body_mass_index
  return bmi

```

Test your code below.  The BMI for a weight of 90.7 kg and height of 182 cm should be approximately 27.38

```python
 body_weight_in_kg= 90.7
 height_in_mts=1.82
 # Computing the body_mass_index
 bmi_output=  bmi(body_weight_in_kg, height_in_mts )
 # Print the result
 print(f"The BMI is: {bmi_output:.2f}")
```

Output:
```
The BMI is: 27.38
```

```python
assert round(bmi(90.7, 1.82), 2) == 27.38
assert round(bmi(91, 1.80), 2) == 28.09
assert round(bmi(70, 1.70), 2) == 24.22
```

```python
bmi(90.7, 1.82)
```

Output:
```
27.38195870064002
```

```python
# Print the result
print(bmi(90.7, 1.82))
```

Output:
```
27.38195870064002
```

### 5.2 Weight from BMI

For the next operation, we're going to calculate Weight in kilograms, given someone's BMI and their height in meters.

This calculation will be:  $ weight\,\mathrm{kg} = {bmi}\times{(height\,\mathrm{m})^2} $

```python
def weight(bmi, height_m):

    ### BEGIN SOLUTION
    w = -1
    ### END SOLUTION

    return w
```

```python
def weight(body_mass_index,height_in_mts):
# Calculate  weight using the formula:  weight= body_mass_index*(height_in_mts**2)
     weight=  body_mass_index*(height_in_mts ** 2)
    # Return to weight

     return weight

```

```python
# Lets assume body_mass_index of a person is 27.38 and height_in_mts is 1.82
body_mass_index= 27.38
height_in_mts=1.82
# Computing weight_output
weight_output=  body_weight_in_kg(body_mass_index,height_in_mts)
# Print the result
print(f"Weight of the person is:{weight_output:.2f} kgs")


```

Output:
```
Weight of the person is:90.69 kgs
```

```python
assert round(weight(27.38, 1.82),2) == 90.69
assert round(weight(28.09, 1.80),2) == 91.01
assert round(weight(24.22, 1.70),2) == 70.00
```

```python
# Lets assume body_mass_index is 27.38 and height_in_mts is 1.82
# Print the result
print(weight(27.38, 1.82))
```

Output:
```
90.693512
```

### 5.3 Height from BMI


For the next operation, we're going to calculate height in meters, given someone's BMI and their weight in kilograms.

This calculation will be:  $ {height\,\mathrm{m}} = \sqrt{\frac{weight\,\mathrm{kg}}{bmi}} $

For this equation, you'll notice that we need to calculate a square root.  Python has a function for this of course, `sqrt`, but that function is not part of the base Python language. Some Python features have to be imported.  The first line of code below imports the `sqrt` function from the [math](https://docs.python.org/3/library/math.html) library.


```python
from math import sqrt
def height(bmi, weight_kg):

    ### BEGIN SOLUTION
    h = -1
    ### END SOLUTION

    return h
```

```python
from math import sqrt
def height(body_mass_index,body_weight_in_kg):
  # Calculate the height using the formula: height = math.sqrt(body_weight_in_kg /body_mass_index )
  height = sqrt(body_weight_in_kg /body_mass_index )
  # Return to height
  return height

```

```python
# Lets assume body_mass_index of a person is 27.38 and body_weight_in_kg is 90.69
body_mass_index= 27.38
body_weight_in_kg= 90.69
# Computing height_output
height_output= height(body_mass_index,body_weight_in_kg)
# Print the result
print(f"Height of the person is:{height_output:.2f} mts")
```

Output:
```
Height of the person is:1.82 mts
```

```python
assert round(height(27.38, 90.69),2) == 1.82
assert round(height(28.09, 91.01),2) == 1.80
assert round(height(24.22, 70.00),2) == 1.70
```

```python
# Lets assume body_mass_index is 27.38 and body_weight_in_kg is 90.69
# Print the result
print(height(27.38, 90.69))
```

Output:
```
1.8199647609749618
```

### 5.4 BMI using Feet, Inches, and Pounds

Now, let's try to to create a function that knows how to first convert between imperial units (US inches and pounds) and SI units before calculating the bmi.  For these calculations use:

* $ 1\,\mathrm{inch} = 2.54\,\mathrm{cm} $
* $ 1\,\mathrm{lb} = 0.45454545\,\mathrm{kg} $


```python
def bmi_us(weight_lb, height_ft, height_in):

    ### BEGIN SOLUTION
    answer = -1
    ### END SOLUTION

    return answer
```

```python
def bmi_us(body_weight_in_lb, height_in_ft, height_in_inches):
  # Convert weight from pounds to kilograms
    body_weight_in_kg = body_weight_in_lb * 0.45454545
  # Convert height from feet and inches to meters
    total_inches = height_in_ft * 12 + height_in_inches
    height_in_mts = total_inches * 0.0254
  # Calculate BMI using the formula: weight / (height**2)
    bmi_output = body_weight_in_kg / (height_in_mts ** 2)
   # Return the output BMI
    return bmi_output

```

```python
# Lets assume
body_weight_in_lb=250
height_in_ft=6
height_in_inches=4
# Computing the body mass index
bmi_output= bmi_us(body_weight_in_lb, height_in_ft, height_in_inches)
# Print the result
print(f"The BMI is:{bmi_output:.2f}")
```

Output:
```
The BMI is:30.49
```

```python
assert round(bmi_us(250,6,4),2) == 30.49
assert round(bmi_us(120,5,0),2) == 23.48
assert round(bmi_us(130,5,10),2) == 18.69
```

```python
# Print the result
print(bmi_us(250, 6, 4))
```

Output:
```
30.494583474450884
```

### 5.5 Target BMI
Assuming that someone doesn't change height over time, what is the approximate number of days it will take for someone to get from the BMI you found in #3 above to a BMI of 24.9 (typically regarded as being in a healthy range) while losing weight at a rate of 0.75 kg per week.

As needed, use variables and multiple calculation steps to make your thought process visible in the code. For instance, rather than using the notebook as a calculator, use variables for intermediate answers.

Again, you will receive points both for correct answers as well as for style.

In addition to the bmi, height, and weight equations we already defined above, here is some of the algebra to help you along:

$ height\,\mathrm{m} = \frac{height\,\mathrm{cm}}{100} $

$ days = {weeks} x {7} $

$ weeks\,required\,to\,change\,weight = \frac{current\,weight - target\,weight}{weight\,change\,per\,week} $

<!-- grade_id: weight-answer, points: 0, grade: false, solution: true -->
```python
def lose_weight(weight, height, kg_per_week, target_bmi):

    # Note that the variables you need area already named for you:
    # height is the height in cm
    # weight is the weight in kg
    # kg_per_week is how much the person wants to lose per week
    # target_bmi is the intended target bmi

    ### BEGIN SOLUTION
    days = -1
    ### END SOLUTION

    return days
```

```python
def lose_weight(weight, height, kg_per_week, target_bmi):
  # converting height from centimeters to meters
  height_in_mts=height/100
  # Calculating the current body mass index
  current_BMI=weight/(height_in_mts**2)
  # calculating the target weight to be reached based on the target BMI
  target_weight=target_bmi*(height_in_mts**2)
  # calculating the weight needed to lose
  weight_to_be_lost = weight - target_weight
  # calculating the number of weeks needed to change the weight
  weeks_required= weight_to_be_lost/kg_per_week
  # calculating the total number of days required
  days=weeks_required*7
  return days



```

```python
weight=90.7
height=182
kg_per_week=0.75
target_bmi=24.9
# calculating number of days required to reach the target BMI
days_required=lose_weight(weight, height, kg_per_week, target_bmi)
print(f"The approximate number of days it will take for someone to reach the target BMI:{days_required}")
```

Output:
```
The approximate number of days it will take for someone to reach the target BMI:76.73157333333342
```

Test your code below.  It should return approximately 76.73 days for the given parameters above.

<!-- grade_id: weight-tests, points: 5, grade: true, solution: false -->
```python
assert round(lose_weight(90.7, 182, 0.75, 24.9),2) == 76.73
assert round(lose_weight(91, 180, 0.75, 24.9),2) == 96.36
assert round(lose_weight(70, 170, 0.75, 22),2) == 59.92
assert round(lose_weight(70, 170, 2, 22),2) == 22.47

```

```python
lose_weight(90.7, 182, 0.75, 24.9)
```

Output:
```
76.73157333333342
```

---

### 5.6 Your choice

Every one of us has done projects where some kind of simple mathematical calculators are required.  Think about something recent from your field of expertise or an area of interest. Choose something that you find interesting and write the Python code required to do that calculation.  These examples show the level of complexity you need to have (not very complex).

* $ bmr_{male} = 66 + (13.7 \times weight\,\mathrm{kg}) + (5 \times height\,\mathrm{m}) - (6.8 \times age\,\mathrm{yr}) $

* $ bmr_{female} = 655 + (9.6 \times weight\,\mathrm{kg}) + (1.8 \times height\,\mathrm{cm}) - (4.7 \times age\,\mathrm{yr}) $

* $ Batting\,Average = \frac{Total\,Hits}{Total\,at\,Bats} $

* $ Area\,of\,circle = \pi r^2 $

* $ F = \frac{1}{2} m v^2 $

Then provide the code for how you would use variables and example values to do the calculation.  You do not to wrap this in a Python function. Just include the code and variables as I've shown below with the BMI calculation that we've been using. As the last command in the cell, print out the value of your calculation as shown in my example.


```python
# Calculating the acceleration of an object
def calculate_acceleration(final_velocity, initial_velocity, time):
# Using the formula acceleration=(final_velocity - initial_velocity) / time
 acceleration=(final_velocity - initial_velocity) / time
 return acceleration

# Let's assume a bike accelerates from 45m/s to 95m/s in 7 seconds. What is the acceleration
final_velocity = 95 # in meters per second
initial_velocity = 45 # in meters per second
time = 7  # in seconds

acceleration= calculate_acceleration(final_velocity, initial_velocity, time)
print(f"The acceleration of the bike is: {acceleration:.2f}  meters per second squared")



```

Output:
```
The acceleration of the bike is: 7.14  meters per second squared
```

```python
def my_function(a, b):
    c = a ** b
    return c
```

```python
my_function(1, 8)
```

Output:
```
1
```

```python
my_function(2, 8)
```

Output:
```
256
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week02/week02_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

